Before you submit this notebook, make sure everything runs as expected in the local test cases. 
Please, paste the solution to the designed cell and do not change anything else.

Also, please, leave your first and last names below

```python
FirstName = ""
LastName = ""
```

---

```python
import sys

import numpy as np
import unittest
import time

import collections
import pickle
import io
```

<!-- grade_id: cell-ac8fc52d8a39ccb4, points: 0, grade: false, solution: true -->
```python
import numpy as np

class LossAndDerivatives:
    @staticmethod
    def mse(X, Y, w):
        """
        X : numpy array of shape (`n_observations`, `n_features`)
        Y : numpy array of shape (`n_observations`, `target_dimentionality`) or (`n_observations`,)
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)

        Return : float
            single number with MSE value of linear model (X.dot(w)) with no bias term
            on the selected dataset.
        
        Comment: If Y is two-dimentional, average the error over both dimentions.
        """

        return np.mean((X.dot(w) - Y)**2)

    @staticmethod
    def mae(X, Y, w):
        """
        X : numpy array of shape (`n_observations`, `n_features`)
        Y : numpy array of shape (`n_observations`, `target_dimentionality`) or (`n_observations`,)
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)
                
        Return: float
            single number with MAE value of linear model (X.dot(w)) with no bias term
            on the selected dataset.

        Comment: If Y is two-dimentional, average the error over both dimentions.
        """

        # YOUR CODE HERE    
        return np.mean(abs((X.dot(w) - Y)))

    @staticmethod
    def l2_reg(w):
        """
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)

        Return: float
            single number with sum of squared elements of the weight matrix ( \sum_{ij} w_{ij}^2 )

        Computes the L2 regularization term for the weight matrix w.
        """
        
        # YOUR CODE HERE
        return np.sum(w ** 2)

    @staticmethod
    def l1_reg(w):
        """
        w : numpy array of shape (`n_features`, `target_dimentionality`)

        Return : float
            single number with sum of the absolute values of the weight matrix ( \sum_{ij} |w_{ij}| )
        
        Computes the L1 regularization term for the weight matrix w.
        """

        # YOUR CODE HERE
        return np.sum(abs(w))

    @staticmethod
    def no_reg(w):
        """
        Simply ignores the regularization
        """
        return 0.
    
    @staticmethod
    def mse_derivative(X, Y, w):
        """
        X : numpy array of shape (`n_observations`, `n_features`)
        Y : numpy array of shape (`n_observations`, `target_dimentionality`) or (`n_observations`,)
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)
        
        Return : numpy array of same shape as `w`

        Computes the MSE derivative for linear regression (X.dot(w)) with no bias term
        w.r.t. w weight matrix.
        
        Please mention, that in case `target_dimentionality` > 1 the error is averaged along this
        dimension as well, so you need to consider that fact in derivative implementation.
        """

        # YOUR CODE HERE
        return 2 * (X.T.dot(X.dot(w) - Y)) / (Y.shape[0] * Y.shape[1])

    @staticmethod
    def mae_derivative(X, Y, w):
        """
        X : numpy array of shape (`n_observations`, `n_features`)
        Y : numpy array of shape (`n_observations`, `target_dimentionality`) or (`n_observations`,)
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)
        
        Return : numpy array of same shape as `w`

        Computes the MAE derivative for linear regression (X.dot(w)) with no bias term
        w.r.t. w weight matrix.
        
        Please mention, that in case `target_dimentionality` > 1 the error is averaged along this
        dimension as well, so you need to consider that fact in derivative implementation.
        """

        # YOUR CODE HERE
        return X.T.dot(np.sign(X.dot(w) - Y)) / (Y.shape[0] * Y.shape[1])

    @staticmethod
    def l2_reg_derivative(w):
        """
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)

        Return : numpy array of same shape as `w`

        Computes the L2 regularization term derivative w.r.t. the weight matrix w.
        """

        # YOUR CODE HERE
        return 2 * w

    @staticmethod
    def l1_reg_derivative(w):
        """
        Y : numpy array of shape (`n_observations`, `target_dimentionality`) or (`n_observations`,)
        w : numpy array of shape (`n_features`, `target_dimentionality`) or (`n_features`,)

        Return : numpy array of same shape as `w`

        Computes the L1 regularization term derivative w.r.t. the weight matrix w.
        """

        # YOUR CODE HERE
        return np.sign(w)

    @staticmethod
    def no_reg_derivative(w):
        """
        Simply ignores the derivative
        """
        return np.zeros_like(w)
```

### Test 0: Initialization (0.01 points)

<!-- grade_id: cell-3473b7b6ffd64d07, points: 0.01, grade: true, solution: false -->
```python
# do not change this cell
loss_and_derivatives = LossAndDerivatives


```

### Test 1: MSE derivative (0.24 points)

<!-- grade_id: cell-e3503c286039ec55, points: 0.24, grade: true, solution: false -->
```python
mse_derivative = LossAndDerivatives.mse_derivative(X_ref, y_ref, w_hat)
assert np.allclose(mse_derivative, ref_dict['mse_derivative'], atol=1e-4)
```

### Test 2: MAE derivative (0.25 points)

<!-- grade_id: cell-e2c4124a6f815118, points: 0.25, grade: true, solution: false -->
```python
mae_derivative = LossAndDerivatives.mae_derivative(X_ref, y_ref, w_hat)
assert np.allclose(mae_derivative, ref_dict['mae_derivative'], atol=1e-4)
```

### Test 3: L1 reg derivative (0.15 points)

<!-- grade_id: cell-69473387a23d8dff, points: 0.15, grade: true, solution: false -->
```python
l2_reg_derivative = LossAndDerivatives.l2_reg_derivative(w_hat)
assert np.allclose(l2_reg_derivative, ref_dict['l2_reg_derivative'], atol=1e-4)
```

### Test 4: L1 reg derivative (0.15 points)

<!-- grade_id: cell-3460671ba79fd04a, points: 0.15, grade: true, solution: false -->
```python
l1_reg_derivative = LossAndDerivatives.l1_reg_derivative(w_hat)
assert np.allclose(l1_reg_derivative, ref_dict['l1_reg_derivative'], atol=1e-4)
```

### Test 5: MSE (0.05 points)

<!-- grade_id: cell-5d56bb4222ee1e7c, points: 0.05, grade: true, solution: false -->
```python
mse = LossAndDerivatives.mse(X_ref, y_ref, w_hat)
assert np.allclose(mse, ref_dict['mse'], atol=1e-4)
```

### Test 6: MAE (0.05 points)

<!-- grade_id: cell-8155e619427ac916, points: 0.05, grade: true, solution: false -->
```python
mae = LossAndDerivatives.mae(X_ref, y_ref, w_hat)
assert np.allclose(mae, ref_dict['mae'], atol=1e-4)
```

### Test 7: L2 reg (0.05 points)

<!-- grade_id: cell-c7ccfce79ecbadc3, points: 0.05, grade: true, solution: false -->
```python
l2_reg = LossAndDerivatives.l2_reg(w_hat)
assert np.allclose(l2_reg, ref_dict['l2_reg'], atol=1e-4)
```

### Test 8: L1 reg (0.05 points)

<!-- grade_id: cell-1e1aa415f6d381f7, points: 0.05, grade: true, solution: false -->
```python
l1_reg = LossAndDerivatives.l1_reg(w_hat)
assert np.allclose(l1_reg, ref_dict['l1_reg'], atol=1e-4)   
```